Diverging roles of TRPV1 and TRPM2 in warm-temperature detection

Author response:

The following is the authors’ response to the previous reviews.

Public Reviews:

Reviewer # 1 (Public Review):

Summary:

The authors use an innovative behavior assay (chamber preference test) and standard calcium imaging experiments on cultured dorsal root ganglion (DRG) neurons to evaluate the consequences of global knockout of TRPV1 and TRPM2, and overexpression of TRPV1, on warmth detection. They find a profound effect of TRPM2 elimination in the behavioral assay, whereas the elimination of TRPV1 has the largest effect on the neuronal responses. These findings are very important, as there is substantial ongoing discussion in the field regarding the contribution of TRP channels to different aspects of thermosensation.

Strengths:

The chamber preference test is an important innovation compared to the standard two-plate test, as it depends on thermal information sampled from the entire skin, as opposed to only the plantar side of the paws. With this assay, and the detailed analysis, the authors provide strong supporting evidence for a role of TRPM2 in warmth avoidance. The conceptual framework using the Drift Diffusion Model provides a first glimpse of how this decision of a mouse to change between temperatures can be interpreted and may form the basis for further analysis of thermosensory behavior.

Weaknesses:

The authors juxtapose these behavioral data with calcium imaging data using isolated DRG neurons. As the authors acknowledge, it remains unclear whether the clear behavioral effect seen in the TRPM2 knockout animals is directly related to TRPM2 functioning as a warmth sensor in sensory neurons. The effects of the TRPM2 KO on the proportion of warmth sensing neurons are very subtle, and TRPM2 may also play a role in the behavioral assay through its expression in thermoregulatory processes in the brain. Future behavioral experiments on sensory-neuron specific TRPM2 knockout animals will be required to clarify this important point.

Reviewer # 1 (Recommendations for the authors):

(1) I have no further suggestions for the authors, and congratulate them with their excellent study.

For the authors information, ref. 42 does contain behavioral data from both male (Fig. 4 and Extended Figure 7) and female (Extended Figure 8) mice.

We thank the referee for pointing out that both males and female mice were tested in the Vandewauw et al. 2018 study. We deliberated whether to include this at the appropriate section of our manuscript (“Limitations of the Study”). But since Vandewauw et al. assessed noxious heat temperatures and we here assess innocuous warmth temperature, we felt that this reference would not add to the clarification whether there are sex differences in Trp channelbased warmth temperature sensing. In particular, we did not want to “use” the argument and to suggest that there are no sex temperature differences in the warmth range just because Vandewauw et al. did not observe major sex differences in the noxious temperature range.

Reviewer #3 (Public Review):

Summary and strengths:

In the manuscript, Abd El Hay et al investigate the role of thermally sensitive ion channels TRPM2 and TRPV1 in warm preference and their dynamic response features to thermal stimulation. They develop a novel thermal preference task, where both the floor and air temperature are controlled, and conclude that mice likely integrate floor with air temperature to form a thermal preference. They go on to use knockout mice and show that TRPM2-/- mice play a role in the avoidance of warmer temperatures. Using a new approach for culturing DRG neurons they show the involvement of both channels in warm responsiveness and dynamics. This is an interesting study with novel methods that generate important new information on the different roles of TRPV1 and TRPM2 on thermal behavior.

Comments on revisions:

Thanks to the authors for addressing all the points raised. They now include more details about the classifier, better place their work in context of the literature, corrected the FOVs, and explained the model a bit further. The new analysis in Figure 2 has thrown up some surprising results about cellular responses that seem to reduce the connection between the cellular and behavioral data and there are a few things to address because of this:

(1) TRPM2 deficient responses: The differences in the proportion of TRPM2 deficient responders compared to WT are only observed at one amplitude (39C), and even at this amplitude the effect is subtle. Most surprisingly, TRPM2 deficient cells have an enhanced response to warm compared to WT mice to 33C, but the same response amplitude as WT at 36C and 39C. The authors discuss why this disconnect might be the case, but together with the lack of differences between WT and TRPM2 deficient mice in Fig 3, the data seem in good agreement with ref 7 that there is little effect of TRPM2 on DRG responses to warm in contrast to a larger effect of TRPV1. This doesn't take away from the fact there is a behavioral phenotype in the TRPM2 deficient mice, but the impact of TRPM2 on DRG cellular warm responses is weak and the authors should tone down or remove statements about the strength of TRPM2's impact throughout the manuscript, for example:

"Trpv1 and Trpm2 knockouts have decreased proportions of WSNs."

"this is the first cellular evidence for the involvement of TRPM2 on the response of DRG sensory neurons to warm-temperature stimuli"

"we demonstrate that TRPV1 and TRPM2 channels contribute differently to temperature detection, supported by behavioural and cellular data"

"TRPV1 and TRPM2 affect the abundance of WSNs, with TRPV1 mediating the rapid, dynamic response to warmth and TRPM2 affecting the population response of WSNs."

"Lack of TRPV1 or TRPM2 led to a significant reduction in the proportion of WSNs, compared to wildtype cultures".

We agree with the referee that the somewhat surprising result of the subtle phenotype in Trpm2 knock-out DRG culture experiments, that became detectable in the course of the new analysis, was overemphasized in the previous version of the manuscript. Per suggestion, we have toned down or removed the statements in the revised manuscript (for the referee to find those changes easily, they are indicated in “track-changes mode” in the submitted document).

(2) The new analysis also shows that the removal of TRPV1 leads to cellular responses with smaller responses at low stimulus levels but larger responses with longer latencies at higher stimulus levels. Authors should discuss this further and how it fits with the behavioral data.

Because these changes shown in Fig. 2E are also subtle (similar to the cellular Trpm2 phenotype discussed above), and because both the “% Responders” (Fig 2.D) and The AUC analysis (Fig. 2F) show a reduction in Trpv1 knock out cultures ––both, at lower and at higher stimulus levels–– we did not want to overstate this difference too much and therefore did not further discuss this aspect in the context of the behavioral differences observed in the Trpv1 knock-out animals.

(3) Analysis clarification: authors state that TRPM2 deficient WSNs show "Their response to the second and third stimulus, however, are similar to wildtype WSNs, suggesting that tuning of the response magnitude to different warmth stimuli is degraded in Trpm2-/- animals." but is there a graded response in WT mice? It looks like there is in terms of the %responders but not in terms of response amplitude or AUC. Authors could show stats on the figure showing differences in response amplitude/AUC/responders% to different stimulus amplitudes within the WT group.

We have added the statistics in the main text, you find them on page 7 (also in “track changes mode”).

(4) New discussion point: sex differences are "similar to what has been shown for an operant-based thermal choice assay (11,56)", but in their rebuttal, they mention that ref 11 did not report sex differences. 56 does. Check this.

Thank you for pointing out this mishap. We have now corrected this in the “Limitations of the study” section of the discussion and have removed the Paricio-Montesions et al study from that section and slightly revised the text (see “track-changes” on page 16).

(5) The authors added in new text about the drift diffusion model in the results, however it's still not completely clear whether the "noise" is due to a perceptual deficit or some other underlying cause. Perhaps authors could discuss this further in the discussion.

We have now included more discussion concerning this (page 14):

“However, the increased noise in the drift-di3usion model points to a less reliable temperature detection mechanism. Although noise in drift di3usion models can encompass various sources of variability—ranging from peripheral sensory processing to central mechanisms like attention or motor initiation—the most parsimonious interpretation in our study aligns with a perceptual deficit, given the altered temperatureresponsive neuronal populations we observed. This implies that, despite the substantial loss of WSNs, the remaining neuronal population provides su3icient information for the detection of warmer temperatures, albeit with reduced precision”

Within the limits of the data that is available, we hope the referee agrees with us that we have now adequately discussed this aspect; we feel that any further discussion would be too speculative.

Changes from the earlier version to the revised version

Abstract

The accurate perception of innocuous temperatures, particularly those experienced as pleasantly warm, is essential for achieving thermal comfort and maintaining thermoregulatory balance. Warm-sensitive neurons (WSN) innervating the skin play a central role in non-painful warmth detection. The TRP ion channels TRPV1 and TRPM2 have been suggested as sensors of warm temperature in WSNs. However, the precise contribution of these channels to the process of warmth detection is not fully understood.

Removed in the revised version: A significant challenge in analysing WSNs lies in their scarcity: fewer than 10% of sensory neurons in the rodent dorsal root ganglion (DRG) respond to innocuous warm temperatures. In this study, we examined >20,000 cultured mouse DRG neurons using calcium imaging and discovered distinct contributions of TRPV1 and TRPM2 to warm-temperature sensitivity. TRPV1 and TRPM2 affect the abundance of WSNs, with TRPV1 mediating the rapid, dynamic response to warmth and TRPM2 affecting the population response of WSNs.

Earlier version: By carefully tracking animal movement in a whole-body thermal preference paradigm, we observe that these cellular differences correlate with nuanced thermal behaviours. Utilizing a driftdiffusion model to quantitatively analyse the decision-making process of animals exposed to different environmental temperatures, we found that: TRPV1 deletion primarily impairs the precision of evidence accumulation, whereas TRPM2 deletion significantly increases the total duration of exposure to warmer environments that are avoided by wildtype mice.

Revised version: A significant challenge in analysing WSNs lies in their scarcity: fewer than 10 % of sensory neurons in the rodent dorsal root ganglion (DRG) respond to innocuous warm temperatures. In this study, we examined >20,000 cultured mouse DRG neurons using calcium imaging and discovered distinct contributions of TRPV1 and TRPM2 to warm-temperature sensitivity. TRPV1, and to a lesser extent TRPM2, affect the abundance of WSNs, with TRPV1 mediating the rapid, dynamic response to warmth and TRPM2 subtly affecting the population response of WSNs. By carefully tracking animal movement in a whole-body thermal preference paradigm, we observe that these cellular differences correlate with nuanced thermal behaviours. Utilizing a drift-diffusion model to quantitatively analyse the decision-making process of animals exposed to different environmental temperatures, we found that: TRPV1 deletion primarily impairs the precision of evidence accumulation, whereas TRPM2 deletion significantly increases the total duration of exposure to warmer environments that are avoided by wildtype mice.

Our findings provide valuable insights into the distinct molecular responses to warmth stimuli, and underpin the subtle aspects of thermal decision-making when encountering minor temperature variations.

Introduction

The detection of temperature, and the related behavioural responses, are an integral part of our sensory interaction with the outside world. Thus, it is not surprising that temperature detection was one of the first sensory modalities to be studied in contemporary neuroscience (1). Early studies concentrated on the characterization of temperature-specific sensory fibres, covering the range from noxious cold, through innocuous cold and warm, to noxious heat (2). However, the molecular mechanism by which temperature activates these fibres remained elusive for decades.

A major breakthrough in the field of somatosensory research was the identification of temperature-sensitive ion channels that belong to the transient-receptor potential (TRP) super family as the molecular sensors responsible for the detection of noxious cold and heat in sensory neurons (3–5). However, the detection of temperatures in-between noxious cold and heat (25 °C to 43 °C), which are often perceived as non-painful, is incompletely understood. This innocuous temperature range also contains the so-called thermoneutral point (TNP); an ambient temperature (29 °C to 33 °C) at which mice do not exert additional energy to maintain their body temperature (6). This makes the innocuous temperature range crucial for thermoregulation and the animal’s subsequent thermal or comfort choice. Recent studies began to uncover the mechanisms behind innocuous warm-temperature detection, thereby mainly converging on three candidate cation channels, namely TRPV1, TRPM2 and TRPM8 (7–11). Interestingly, the evidence for the involvement of TRPV1 and TRPM2 channels is seemingly contradictory.

TRPV1 is traditionally associated with the response to noxious temperature stimuli (> 42 °C), with the ability to become sensitive to lower temperatures in inflammatory contexts (3, 12–14). However, in vivo calcium imaging of trigeminal sensory neurons in animals lacking TRPV1 showed a complete absence of responses to warm stimuli applied to the oral cavity of mice, while responses to hot temperatures were unchanged. Furthermore, acute inhibition of TRPV1 in animals trained to discriminate innocuous cold from warmth through a nose port led to a reduction in their performance (7). A similar result was described in another operant behavior task where animals had to report a warming stimulus applied to the paws (11). These results stand in contrast to other studies that showed no involvement of TRPV1 in thermal preference across the warm-temperature range (15–17).

Earlier version: TRPM2 has reported ex vivo activation temperatures between 35 °C and >40 °C, depending on the cellular context, and was first described as a physiological temperature sensor in pancreatic islet cells (18, 19). Calcium imaging of DRG cultures from animals lacking TRPM2 showed a reduction in the proportion of warm- and heat-responsive neurons in comparison to wildtype sensory neuron cultures (8). Interestingly, Trpm2−/− animals are unable to differentiate temperatures across the innocuous warm range in thermal preference tasks (8, 20), while their ability to avoid noxious temperatures is not affected. In-terestingly, and similar to TRPV1, Trpm2-deficient animals trained to report warming of their paws were less sensitive than wildtype animals (11).

Revised version: TRPM2 has reported ex vivo activation temperatures between 35 °C and >40 °C, depending on the cellular context, and was first described as a physiological temperature sensor in pancreatic islet cells (18, 19). Calcium imaging of DRG cultures from animals lacking TRPM2 showed a reduction in the proportion of warm- and heat-responsive neurons in comparison to wildtype sensory neuron cultures (8). Interestingly, Trpm2−/− animals are unable to differentiate temperatures across the innocuous warm range in thermal preference tasks (8, 20), while their ability to avoid noxious temperatures is not affected. Interestingly, and similar to TRPV1, Trpm2-deficient animals trained to report warming of their paws were less sensitive than wildtype animals (11).

In summary, both for TRPV1 and TRPM2, there is an apparent disconnect between the observations in cellular assays and the behavioural tasks assessing temperature detection. In vivo calcium imaging coupled with warm-temperature stimuli only shows a relevance of TRPV1, but not TRPM2 in the innocuous temperature range (7). Contrary to that, the lack of TRPM2, but not TRPV1, more consistently affects warm-temperature detection in assays of temperature preference (8, 15–17), with both affecting the animals’ temperature perception in operant behavioural assays, albeit only subtly (7, 11).

One main challenge for the analysis of WSNs, neurons that respond to innocuous temperature stimuli between 25 °C and 43 °C, is the low abundance of this neuronal population, which represents about 3 % to 10 % of sensory neurons in rodents (8, 21). Using in-depth functional analysis of 1000s of sensory neurons from multiple animals, we here describe that TRPV1 and TRPM2 are both involved in the detection of innocuous (warm) temperature stimuli. Furthermore, we demonstrate the diverging roles both channels play in warm-temperature detection through a novel thermal preference behaviour assay.

Results

The thermal chamber preference test allows precise discrimination of subtle temperature differences in the innocuous range

Mice prefer 31 °C to warmer temperatures when ambient and floor temperature are controlled

Earlier version: The ability to avoid uncom-fortable environmental temperatures or move toward pleasant thermal conditions is fundamental to sustaining life. This process is known as behavioural thermoregulation and can be found in most groups of animals (22). The preference of rodents for temperature is traditionally assessed via paradigms that challenge the animals with differing floor temperatures (23–25). This leads to preference development that is based on temperature detection through glabrous skin, such as the paws, tail, and nose. Animals, however, are capable of in-tegrating temperature from glabrous and non-glabrous skin (26). With the aim to probe behavioural thermoregulation in a more holistic context, we developed a thermal discrimination assay where both the floor and ambient temperature are controlled, termed the thermal chamber preference test (CPT) (Figure 1A-D, Figure S1). When presented with warm ambient temperatures (34 °C and 38 °C) and 31 °C as control temperature, mice significantly preferred the 31 °C chamber over the warmer chambers (Figure 1E and F). Compared to the classic two-plate preference test (TPT), wildtype animals developed a stronger preference for the 31 °C side in the CPT (Figure S1I). Additionally, animals showed a clear preference for 31 °C when given 34 °C as an option in the CPT. This is not observed in the classic TPT (Figure S1I). This observation suggests that more subtle ambient temperature differences, relating to comfort and thermoregulation, are more faithfully assessed in the CPT assay.

Revised version: The ability to avoid uncomfortable environmental temperatures or move toward pleasant thermal conditions is fundamental to sustaining life. This process is known as behavioural thermoregulation and can be found in most groups of animals (22). The preference of rodents for temperature is traditionally assessed via paradigms that challenge the animals with differing floor temperatures (23–25). This leads to preference development that is based on temperature detection through glabrous skin, such as the paws, tail, and nose. Animals, however, are capable of integrating temperature from glabrous and non-glabrous skin (26). With the aim to probe behavioural thermoregulation in a more holistic context, we developed a thermal discrimination assay where both the floor and ambient temperature are controlled, termed the thermal chamber preference test (CPT) (Figure 1A-D, Figure S1). When presented with warm ambient temperatures (34 °C and 38 °C) and 31 °C as control temperature, mice significantly preferred the 31 °C chamber over the warmer chambers (Figure 1E and F). Compared to the classic two-plate preference test (TPT), wildtype animals developed a stronger preference for the 31 °C side in the CPT (Figure S1I). Additionally, animals showed a clear preference for 31 °C when given 34 °C as an option in the CPT. This is not observed in the classic TPT (Figure S1I). This observation suggests that more subtle ambient temperature differences, relating to comfort and thermoregulation, are more faithfully assessed in the CPT assay.

A novel ambient temperature preference test.

A Schematic of the chamber preference test from the top. Grey outlines the outer enclosure and the dashed line the internal cage. Peltier elements (grey oval shapes) were combined with fans for precise control of the temperature. See Figure S1 for a more detailed view. B A representative image of an animal exploring the chambers. Coloured dots represent the tracked keypoints on the animal and reference points in the enclosure. C Tracking of an example animal for 30 minutes at 31 °C (right chamber) and 34 °C (left chamber). D Density maps of the x-position of the animal in C over 30 minutes; binned in three-minute long intervals. Dashed lines represent the tunnel connecting both chambers. E Density maps as in D with one-minute bins of all animals from wildtype (n = 48), Trpv1−/− (n = 15) and Trpm2−/− (n = 28) genotypes. F Proportion of time spent in the test chamber for animals shown in E over time, binned in three-minute long intervals. ANOVA over genotype (34 °C: F(2,70)=7.30, p=0.001; 38 °C: F(2,76)=7.84, p <0.001), time (34 °C: F(5.63,394.35)=2.51, p <0.05; 38 °C: F(3.86,293.72)=11.07, p <0.001), and their interaction (34 °C: F(11.27,394.35)=1.98, p <0.05). Results of post hoc multiple comparison by timepoint against wildtype are indicated. G Exemplary behaviour of the animal in C and D over the first 120 seconds of the experiment, highlighting the visit frequency and duration of time spent in each chamber. The dashed line represents the tunnel connecting the chambers. H Overview of the frequency and length of the visits to the test chamber for 15 randomly sampled animals per genotype, shown in E. Each visit is coloured by the log2 of its length to highlight varying visit lengths. I averaged and smoothed visit lengths in a 3-minute rolling window with a 1-minute lag. Linear mixed model over genotype and time with random effects across animals. J averaged and smooothed number of crosses in a 3-minute rolling window. Cox regression over genotype (34 °C: χ2 (2)=49.67, p <0.001; 38 °C: χ2(2)=55.74, p <0.0001). I and J Results of post hoc multiple comparison against wildtype are indicated. * (p <0.05), ** (p <0.01), *** (p <0.001), **** (p <0.0001).

TRPM2 is necessary for establishing a preference in the warm-temperature range

Earlier version: Previous studies assessing warmtemperature detection using the TPT with animals lacking TRPV1 or TRPM2 showed that Trpm2−/− animals failed to differentiate between 31 °C and 38 °C (reproduced in Figure S1) while Trpv1−/− animals, similar to wildtypes, preferred the 31 °C side (8, 16). Similarly, when using the newly developed CPT, Trpv1−/− animals showed a similar temperature preference to wildtypes while Trpm2−/− animals failed to develop a preference for the thermoneutral (31 °C) side (Figure 1E-F), without affecting their preference at 25 °C (Figure S2A and C). In addition to the previously described phenotype at 38 °C, Trpm2−/− animals were also unable to discriminate 34 °C from 31 °C, emphasizing the relevance of TRPM2 at milder warm temperatures (Figure 1E-F, and Figure S1). Notably, the phenotype of Trpm2−/− animals was similar to that of animals lacking most, if not all, peripheral (heat and cold) thermosensors (Trpv1Abl, (27), Figure S1J-L). These results confirm previous data demonstrating the requirement for TRPM2 rather than TRPV1, in preference development in the warm-temperature range (8, 16).

Revised version: Previous studies assessing warm-temperature detection using the TPT with animals lacking TRPV1 or TRPM2 showed that Trpm2−/− animals failed to differentiate between 31 °C and 38 °C (reproduced in Figure S1) while Trpv1−/− animals, similar to wildtypes, preferred the 31 °C side (8, 16). Similarly, when using the newly developed CPT, Trpv1−/− animals showed a similar temperature preference to wildtypes while Trpm2−/− animals failed to develop a preference for the thermoneutral (31 °C) side (Figure 1E-F), without affecting their preference at 25 °C (Figure S2A and C). In addition to the previously described phenotype at 38 °C, Trpm2−/− animals were also unable to discriminate 34 °C from 31 °C, emphasizing the relevance of TRPM2 at milder warm temperatures (Figure 1E-F, and Figure S1). Notably, the phenotype of Trpm2−/− animals was similar to that of animals lacking most, if not all, peripheral (heat and cold) thermosensors (Trpv1Abl, (27), Figure S1J-L). These results confirm previous data demonstrating the requirement for TRPM2 rather than TRPV1, in preference development in the warm-temperature range (8, 16).

TRPV1 and TRPM2 affect different aspects of warm-temperature detection

Traditionally, analyses of temperature preference are limited to reporting the proportion of time an animal spent at the test temperature, without assessing more finegrained thermal preference behaviour, such as the sequence of chamber crossings and intermittent pauses (visit lengths, (23– 25)). We observed that in the CPT, animals cross from one chamber to the other, probing the chamber, before crossing back (Figure 1G). We quantified the number of crossings and the lengths of these episodes throughout the experiments (Figure 1H). At the start of the experiment, mice of all genotypes crossed more often than at the end of the experiment while maintaining similar durations of their visits to the warmer chamber (Figure 1H-J). Trpv1−/− animals showed a significantly higher crossing rate compared to wildtype animals at both 34 °C and 38 °C (Figure 1H and J). Trpm2−/− animals, on the other hand, had significantly longer visits to the warmer chamber compared to wildtype animals, while having either similar (at 38 °C) or a reduced crossing rate (at 34 °C) (Figure 1H-I).

Notably, animals lacking TRPM8, a cold-sensitive TRP-channel that was shown to be critical for learning to report warming in an operant behavioural assay (11), showed similar preference behaviour to wildtype animals across the warm-temperature range, but increased visit lengths at 25 °C and fewer crosses across all tested temperatures (Figure S3A-E).

Together, these observations suggest that, contrary to previous findings using the TPT, both TRPV1 and TRPM2 contribute to the animals’ ability to detect warm temperatures and to drive associated thermal preference behaviours, albeit through different behaviors.

Earlier version: Trpv1 and Trpm2 knockouts have decreased proportions of WSNs

Revised version: Trpv1 knockout mice have decreased proportions of WSNs

A small subpopulation of cultured primary sensory neurons responds to warm temperatures

Earlier version: Trpv1 and Trpm2 are highly expressed in peripheral sensory neurons that reside bilaterally in so-called dorsal root ganglia (DRG) along the spinal cord. To assess the individual contribution of both channels to ambient warm-temperature detection, and to account for the integration of temperature across the whole body of the animal, we cultured primary DRG neurons pooled from across the whole length of the spine.

Revised version: Trpv1 is highly expressed in peripheral sensory neurons that reside bilaterally in so-called dorsal root ganglia (DRG) along the spinal cord. To assess the individual contribution of Trpv1 and Trpm2 channels to ambient warm-temperature detection, and to account for the integration of temperature across the whole body of the animal, we cultured primary DRG neurons pooled from across the whole length of the spine.

Historically, experiments studying temperature responses in sensory neurons are performed with DRG neurons cultured for a few hours to overnight (28). However, these cultures did not reflect the distribution of WSNs described from in vivo studies, with 26 ± 9 % of all cells responding to warm temperatures, contrary to 3 % to 10 % of WSNs observed in vivo (Figure S4C-E, (7, 21)). We speculated that this expansion in the proportion of WSNs might reflect a post-injury state in which heat-sensitive neurons become sensitised to lower thermal stimuli (7, 29). Since sensory neuron dissociation resembles an axotomy which activates injury-related pathways (30–35), we extended the commonly used DRG primary culture protocol to three days, to allow the cells to recover from the procedure. Three-day cultures harboured approximately 6 ± 3 % WSNs, compared to overnight culture’s 26 ± 9 % (Figure S4E). Furthermore, three-day cultures showed a reduced calcium inflow upon temperature stimulation, and an improved capacity to return to baseline calcium levels upon termination of the stimulus, indicative of recovery from a of post-injury state of sensory neurons (Figure S4E-G, (32, 33)). These observations are in line with data collected from in vivo calcium imaging preparations of both dorsal root and trigeminal ganglion cells in response to warm temperatures (7, 21) and suggest that three-day cultures – rather than acute/short-term preparations – more accurately reflect the functional properties and abundance of warm-responsive sensory neurons that are found in behaving animals. However, whether 3-day cultures resemble native sensory neurons more closely than acute cultures in terms of their (transcriptional) identity is currently unknown.

Trpv1 and Trpm2 deletion reduces the proportion of warmth responders

Earlier version: To investigate the effects of TRPV1 and TRPM2 loss on the response of sensory neurons to warm-temperature stimulation, we applied the same stimulation protocol to cultures from Trpv1−/− and Trpm2−/− animals (Figure 2). Lack of TRPV1 or TRPM2 led to a significant reduction in the proportion of WSNs, compared to wildtype cultures (Figure 2D, Wildtype : 6.4 ± 2.4%; Trpv1−/−: 2.3 ± 1.3%; Trpm2−/−: 4.7 ± 1.4%). Cultures from Trpv1−/− animals had reduced proportions of responders across the whole range of warm-temperature stimuli (Figure 2D), but showed similar proportions of heat responders (neurons responding to T≥ 43 °C) when compared to cultures obtained from wildtype animals (Figure S5D). In contrast to a previous study describing WSNs in vivo (7), our Trpv1−/− cultures did not show a complete absence of response to warm temperatures, with some cells in the Trpv1−/− cultures retaining their ability to respond to warm stimuli (Figure 2B and D). Lack of TRPM2, on the other hand, affected the proportions of responders across both the warm and the hot temperature range, albeit only subtly (Figure 2D and Figure S5D).

Revised version: To investigate the effects of TRPV1 and TRPM2 loss on the response of sensory neurons to warm-temperature stimulation, we applied the same stimulation protocol to cultures from Trpv1−/− and Trpm2−/− animals (Figure 2). Lack of TRPV1 or TRPM2 led to a significant reduction in the proportion of WSNs, compared to wildtype cultures, albeit the deletion of Trpm2 had only a fairly small effect (Figure 2D, Wildtype : 6.4 ± 2.4 %; Trpv1−/−: 2.3 ± 1.3 %; Trpm2−/−: 4.7 ± 1.4 %). Cultures from Trpv1−/− animals had reduced proportions of responders across the whole range of warm-temperature stimuli (Figure 2D), but showed similar proportions of heat responders (neurons responding to T ≥ 43 °C) when compared to cultures obtained from wildtype animals (Figure S5D). In contrast to a previous study describing WSNs in vivo (7), our Trpv1−/− cultures did not show a complete absence of response to warm temperatures, with some cells in the Trpv1−/− cultures retaining their ability to respond to warm stimuli (Figure 2B and D). Lack of TRPM2, on the other hand, affected the proportions of responders across both the warm and the hot temperature range (Figure 2D and Figure S5D).

Earlier version: Loss of TRPV1 or TRPM2 leads to a reduction in WSN abundance.

Revised version: TRPV1 or and to a lesser extent TRPM2 leads to a reduction incontribute to WSN abundance in DRG cultures.

A Experimental paradigm of temperature stimulation. Three sequential and increasing temperature stimuli of 25 seconds, with 5 minute inter-stimulus intervals. Traces represent mean temperatures for wildtype, Trpv1−/−, and Trpm2−/− cultures. B Heat map showing representative normalized (ΔF/F0) calcium response of WSNs (60 randomly sampled cells per genotype). C Zoom-in of the mean and standard deviation of the three warm-temperature stimuli shown in A. D The proportions of responders to each stimulus in relation to all imaged neurons from wildtype (5 animals, 18 FOVs, 6374 cells), Trpv1−/− (5 animals, 9 FOVs, 3009 cells), and Trpm2−/− (6 animals, 12 FOVs, 4315 cells). Each dot represents a field of view (FOV). ANOVA over genotype (F(2,36) = 14.24, p <0.001), stimulus (F(1.31, 47.32) = 113.44, p <0.001), and their interaction (F(2.63, 47.32) = 9.18, p <0.001). Results of post hoc multiple comparison against wildtype are indicated. E Average and SEM ΔF/F0 for all responders over the whole stimulus (as shown in B) from wildtype (412 cells), Trpv1−/− (111 cells), and Trpm2−/− (204 cells). F Mean area under the curve (AUC) of (ΔF/F0) from each FOV used in D. Linear mixed model over genotype and stimulus with random effects across animals and FOVs. Pairwise contrasts against wildtype are indicated. * (p <0.05), ** (p <0.01), *** (p <0.001).

Added in the revised version: Notably, WSNs from wildtype animals show an increase in the proportion of responders with increased stimulus intensity (second vs. first p < 0.05, third vs. second p<0.0001). This was not observed in cultures from Trpv1−/− or Trpm2−/−, where the increase was only significant from the second to the third stimulus. (p < 0.001 and p< 0.0001, respectively).

Added in the revised version: Loss of Trpm2 alters the population response profile to warm stimuli

Added in the revised version: Next, we compared the magnitude of the responses of WSNs. Previous studies suggested that WSNs are tuned in a graded-monotonic way, i.e., an increase in temperature leads to an increased response (7, 21). Wildtype and Trpv1−/− WSNs show an increase in response magnitude (ΔF/F0) with increasing temperature stimuli (Figure 2E and F). Surprisingly, WSNs from Trpm2−/− animals, on the other hand, respond with a significantly higher calcium inflow to the lowest temperature stimulus, compared with wildtype WSNs. Their response to the second and third stimulus, however, are similar to wildtype WSNs, suggesting that tuning of the response magnitude to different warmth stimuli might be affected in Trpm2−/− animals. When testing for the graded increase in response amplitude, WSNs from wildtype or Trpv1−/− only animals showed a significant increase comparing the second to the third stimulus (p<0.0001 and p<0.05, respectively). Trpm2−/− cells showed a significant decrease from the first to the second stimulus (p<0.001) followed by a significant increase from the second to third stimulus (p<0.0001).

WSNs vary in their response characteristics

Removed in the revised version: Loss of Trpm2 alters the population response profile to warm stimuli

Removed in the revised version: Next, we compared the magnitude of the responses of WSNs. Previous studies suggested that WSNs are tuned in a graded-monotonic way, i.e., an increase in temperature leads to an increased response (7, 21). Wildtype and Trpv1−/− WSNs show an increase in response magnitude (ΔF/F0) with increasing temperature stimuli (Figure 2E and F). Surprisingly, WSNs from Trpm2−/− animals, on the other hand, respond with a significantly higher calcium inflow to the lowest temperature stimulus, compared with wildtype WSNs. Their response to the second and third stimulus, however, are similar to wildtype WSNs, suggesting that tuning of the response magnitude to different warmth stimuli is degraded in TRPM2 KO animals. Combined, this is the first cellular evidence for the involvement of TRPM2 on the response of DRG sensory neurons to warm-temperature stimuli, given that previous studies had only assessed noxious/painful heat responses (T ≥ 45 °C (8)).

TRPV1 drives the dynamic phase of warm-temperature responses

Consistent with the behavioural data (Figure 1), the absence of TRPV1 or TRPM2 led to population changes in response to warm temperatures in primary sensory neurons. A closer look at the calcium response profiles of individual cells showed that WSNs also vary in when they respond to a temperature stimulus (Figure 2B and Figure 3B). To capture this variability, we computed the point at which each cell started responding to the stimulus (Figure 3B).

Trpv1−/− diminishes the response to dynamic temperature changes.

A Temperature traces from three exemplary imaging sessions. B Individual calcium traces (ΔF/F0) of 10 representative thermosensitive neurons from each genotype in response to the applied stimuli. The position of the dashed line indicates the time when the cells exceeded 10 % of their maximum ΔF/F0 during the stimulus. C Mean and standard deviation of the three warm-temperature stimuli shown in A. The dotted line indicates the separation between the dynamic and static phases, defined by the end of the peak of the smoothed temperature change rate. D Response onset of all recorded WSNs. Each row represents a single FOV (see Figure 2D). Each triangle indicates the time point at which the individual cell responds to the stimulus as shown in B. Dotted line as in C. E Density plot of response time points for each genotype and stimulus. Distributions were compared against wildtype using the Wilcoxon ranked sum test with false-discovery rate post hoc correction. * (p <0.05), **** (p <0.0001)..

Each temperature stimulus we divided into two distinct phases: an initial, dynamic phase, in which the temperature rises rapidly. And a second, static phase, in which the temperature stabilizes (Figure 3A). The majority of WSNs from wildtype animals respond during the rising, dynamic phase of the stimulus (Figure 3D and E). In comparison, WSNs from animals lacking TRPV1 predominantly responded during the static phase of the stimulus, while Trpm2−/− cells did not significantly differ in their response onset from wildtype cells (Figure 3D and E and Figure S6A-C). Additionally, TRPV1-positive cells in wildtype cultures – identified by their response to the TRPV1 activator capsaicin (3) – predominantly responded during the dynamic phase of the stimulus, compared to TRPV1-negative cells (Figure S6D-F). Collectively, these observations suggest that TRPV1, but not TRPM2, is involved in the response to dynamic, fast changes in temperature.

Given the strong reduction of the intracellular calcium dynamics observed in Trpv1−/− DRGs exposed to warm stimuli, we speculate that over-expression of Trpv1 would alter the response dynamics of WSNs, particularly during the rising phase of the stimulus. To test this hypothesis, we made use of a previously described animal model which over-expresses Trpv1 in TRPV1-positive cells (Trpv1OX, (36)). WSNs from Trpv1OX animals showed a significantly higher propensity to respond during the dynamic phase of the stimulus, when compared to wildtype cultures (Figure 4A-E). These results align with our previous observations and further suggest that TRPV1 abundance directly regulates the onset and speed of a temperature response. Notably, DRG cultures from Trpv1OX animals showed nearly double the proportion of WSNs compared to wildtype cultures (Wildtype: 6.8 ± 3.9 %; Trpv1OX: 12.8 ± 0.6 %), which suggests that TRPV1-overexpression reduces the response threshold of WSNs.

High TRPV1 expression levels promote dynamic warm-temperature detection and enhance temperature preference.

A Mean temperatures from all experiments and imaging sessions for wildtype (3 animals, 6 FOVs, 3133 cells) and Trpv1OX (2 animals, 5 FOVs, 3754 cells) cultures. B Examples of normalized (ΔF/F0) calcium responses of WSNs responding to any of the stimuli depicted in A. 42 randomly sampled cells from each genotype. C Mean and standard deviation of the three warm-temperature stimuli applied. The dotted line indicates the separation between the dynamic phase and the static phase (see Figure 3C). D Response initiation of all WSNs imaged from wildtype and Trpv1OX animals. Each row represents an individual imaging session. Each triangle denotes the time point at which the cell responds to the stimulus, as shown in Figure 3B. Dotted line as in C. E Density representation of response time points for each genotype and stimulus. Distributions were compared against wildtype using the Wilcoxon ranked sum test. F Density maps of all wildtype (n = 48) and Trpv1OX (n = 12) animals in the CPT over time. G Proportion of time spent in the test chamber for animals shown in F over time, binned in 3-minute intervals. ANOVA over genotype (34 °C: F(1,51)=6.49, p<0.05), time (34 °C: F(6.12,311.93)=4.49, p <0.001; 38 °C: F(4.71,240.22)=28.27, p <0.001), and their interaction (38 °C: F(4.71,240.22)=4.42, p<0.001). Results of post hoc multiple comparison by timepoint against wildtype are indicated. H Overview of the frequency and length of the visits to the test chamber for all animals shown in F. Each visit is coloured by the log2 of its length to highlight varying visit lengths. I averaged and smoothed visit lengths in a 3-minute rolling window with a 1-minute lag. Linear mixed model over genotype and time with random effects across animals. J averaged and smoothed number of crosses in a 3-minute rolling window. Cox regression over genotype (34 °C: χ2 (2)=17.52, p <0.001; 38 °C: χ2(2)=4.13, p <0.05). * (p <0.05), ** (p <0.01), *** (p <0.001), **** (p <0.0001).

Does the enrichment of cells responding during the dynamic stimulus phase affect the behaviour of the animals in the CPT? Indeed, Trpv1 over-expression led to a significantly stronger avoidance of the 38 °C side in the CPT (Figure 4F and G). Interestingly, Trpv1OX animals crossed significantly less between chambers, compared to wildtype animals, while having similar duration of stays at the test chamber (Figure 4H, I, and J), suggesting that Trpv1OX animals discriminate temperatures more rapidly.

An evidence-accumulation model uncovers differences in evidence accumulation across genotypes

Earlier version: In the previous sections, we have detailed the distinct effects that TRPV1 and TRPM2 exert on the temperature responses of sensory neurons. While these cellular-level findings are illuminating, they present a challenge when it comes to directly relating them to the more complex, multifaceted behaviours observed in our temperature preference assay. To bridge this gap, and extract parameters that could be directly correlated with the neuronal data, we conceptualized the temperature preference assay as a continuous decision-making process (Figure 5A and Figure 5B), allowing the use of established evidence accumulation frameworks. These models have been shown to successfully recapitulate animal and human behaviour in sensory decision tasks involving different modalities (37–39) and have even been directly linked to neural observations (40, 41).

Revised version: In the previous sections, we have detailed the distinct effects that TRPV1 and, to a lesser extent, TRPM2 exert on the temperature responses of sensory neurons. While these cellular-level findings are illuminating, they present a challenge when it comes to directly relating them to the more complex, multifaceted behaviours observed in our temperature preference assay. To bridge this gap, and extract parameters that could be directly correlated with the neuronal data, we conceptualized the temperature preference assay as a continuous decision-making process (Figure 5A and Figure 5B), allowing the use of established evidence accumulation frameworks. These models have been shown to successfully recapitulate animal and human behaviour in sensory decision tasks involving different modalities (37–39) and have even been directly linked to neural observations (40, 41).

Modelling the varying roles of TRPV1 and TRPM2 on warm-temperature detection.

A Two example episodes of an animal inside the CPT, crossing from one chamber to the other (see Figure 1G). Dashed line represents crossing time points between chambers. B Examples of possible evidence accumulation process for the two episodes in A using a Drift Diffusion Model (DDM). C and D Simulations of a drift diffusion process with fixed starting points (z = 0.5) and bound (a = 1) while varying drift rates (v in C) and noise (sv in D). E and F depict the resulting distributions of visit lengths when simulating 1000 trials with the parameters from C and D, respectively. G Distributions of visit lengths at 31 °C vs. 34 °C in wildtype animals. Insets show the estimated parameters for v and sv at each temperature and the Kullback-Leibler (KL) divergence between the model (continuous density line) and the empirical data (histogram). See Figure S7 for all model fits. H Drift v and noise sv estimates for both neutral (31 °C, solid line) and test (25 °C to 38 °C, dashed line) chambers for wildtype animals resulting from hierarchical Markov chain Monte Carlo (MCMC) sampling. I Neutral chamber (31 °C) corrected and wildtype-subtracted estimates of drift and noise for all genotypes. The dashed line represents the wildtype reference. Points indicate individual MCMC samples, and vertical lines the median of each distribution.

Earlier version: In our experimental setup, an animal enters a chamber and begins to accumulate evidence (i.e., it continuously collects and computes spatial temperature information) that drives its decision to stay or leave the chamber (Figure 5A). The resulting time spent in each chamber varies from visit to visit (Figure 1G and Figure 5A), highlighting the need to account for a dynamic and stochastic decision-making process. This stochastic variability could be due to many factors, such as variability between animals, fluctuations in cognitive variables (attention or motivation, for example), different exploration speeds of the mice (which alters the duration the mice are exposed to different temperatures, and which could result in a perceived change in perceptual threshold) or actual changes in perceptual threshold over time. Traditional approaches to analysing these behaviours have focused on final outcomes or mean stay time and thus may have overlooked these more subtle nuances (8, 16, 42, 43), without considering the fluctuating nature of sensory perception and the variability of decisionmaking across visits and between individual animals.

Revised version: In our experimental setup, an animal enters a chamber and begins to accumulate evidence (i.e., it continuously collects and computes spatial temperature information) that drives its decision to stay or leave the chamber (Figure 5A). The resulting time spent in each chamber varies from visit to visit (Figure 1G and Figure 5A), highlighting the need to account for a dynamic and stochastic decision-making process. This stochastic variability could be due to many factors, such as variability between animals, fluctuations in cognitive variables (attention or motivation, for example), different exploration speeds of the mice (which alters the duration the mice are exposed to different temperatures, and which could result in a perceived change in perceptual threshold) or actual changes in perceptual threshold over time. Traditional approaches to analysing these behaviours have focused on final outcomes or mean stay time and thus may have overlooked these more subtle nuances (8, 16, 42, 43), without considering the fluctuating nature of sensory perception and the variability of decision-making across visits and between individual animals.

A drift diffusion model recapitulates the animals’ behaviour in the CPT

To model this decision-making process, we opted for a drift diffusion model (DDM (44)). This model provides a quantitative framework to delve into how animals integrate sensory information over time, leading to their decision to stay or leave a chamber (Figure 5).

Earlier version: When an animal enters a chamber, it starts at a certain point (z) from which the accumulation of sensory information/evidence begins (Figure 5B). Over time, and while the temperature information is integrated, the information is accumulated at a certain rate (drift rate, v) towards a decision point (decision bound, a). Additionally, the drift rate is also allowed to fluctuate (noise, sv), which represents the variability of information accumulation (sensory perception). This process evolves until the decision bound (a) is reached, which prompts the animal to leave the chamber. Figures 5C and D show simulations of the DDM, which allow an intuition into how varying levels of drift and noise (while keeping the starting point and decision bound fixed) alters the distribution of visit lengths throughout an experiment (Figures 5E and F). Higher rates of drift (with fixed noise levels) lead to overall shorter visit lengths (Figure 5E). High noise levels (with fixed drift rates) also lead to shorter visit lengths, albeit with larger variability (Figure 5F).

Revised version: When an animal enters a chamber, it starts at a certain point (z) from which the accumulation of sensory information/evidence begins (Figure 5B). Over time, and while the temperature information is integrated, the information is accumulated at a certain rate (drift rate, v) towards a decision point (decision bound, a). Additionally, the drift rate is also allowed to fluctuate (noise, sv), which represents the variability of information accumulation (sensory perception). This process evolves until the decision bound (a) is reached, which prompts the animal to leave the chamber. Figures 5C and D show simulations of the DDM, which allow an intuition into how varying levels of drift and noise (while keeping the starting point and decision bound fixed) alters the distribution of visit lengths throughout an experiment (Figures 5E and F). Higher rates of drift (with fixed noise levels) lead to overall shorter visit lengths (Figure 5E). High noise levels (with fixed drift rates) also lead to shorter visit lengths, albeit with larger variability (Figure 5F). We fit a model with varying drift rate and noise onto the visit length distributions of wildtype animals recorded for each temperature combination in the CPT (Figure 5G and Figure S7). The model suggests higher drift rates at temperatures below and above 31 °C in wildtype animals (Figure 5H). This means that the sensory information to leave the non-neutral chambers is accumulated faster. The model also suggests an increase in noise in the test chambers, hinting that the evidence accumulation at 31 °C is particularly stable (Figure 5H). Both observations are in line with the development of preference for the 31 °C chamber in the varying temperature conditions.

Removed in the revised version: We fit a model with varying drift rate and noise onto the visit length distributions of wildtype animals recorded for each temperature combination in the CPT (Figure 5G and Figure S7). The model suggests higher drift rates at temperatures below and above 31 °C in wildtype animals (Figure 5H). This means that the sensory information to leave the non-neutral chambers is accumulated faster. The model also suggests an increase in noise in the test chambers, hinting that the evidence accumulation at 31 °C is particularly stable (Figure 5H). Both observations are in line with the development of preference for the 31 °C chamber in the varying temperature conditions.

Varying effects of TRPM2 and TRPV1 on evidence accumulation

We fit the model onto all behavioural data collected in this study (Figure 5I and Figure S7). For animals lacking TRPM2, the model yields a lower drift rate at 34 °C and 38 °C compared to wildtypes (Figure 5I). This suggests that loss of TRPM2 leads to a slower evidence accumulation at warm temperatures, reflecting an overall failure to develop a preference for 31 °C throughout the experiment (Figure 1). In Trpv1−/− animals, on the other hand, we observed higher drift rates as well as higher noise levels in the warmer (34 °C and 38 °C) chambers (Figure 5I). These findings suggest that Trpv1−/− animals accumulate environmental temperature evidence faster than wildtype animals, but the fidelity of the thermal inputs is compromised. We speculate that the balance of these two variables might lead to a similar overall preference development compared to wildtype animals (Figure 1). Interestingly, the over-expression of TRPV1 also leads to an increased drift rate at warm temperatures, albeit with a similar noise level to wildtypes. This combination leads to greater avoidance of 34 °C and 38 °C (as observed in Figure 4).

Notably, all genotypes show similar drift and noise estimates at 25 °C, consistent with their behavioural preference (Figure S2, suggesting that TRPV1 and TRPM2 mainly control responses to warm temperatures (Figure 5I). In summary, we find that the DDM successfully parametrizes the behavioural data obtained from the CPT. Furthermore, the model allowed a more in-depth insight into how the loss of either TRPV1 or TRPM2 differentially alters the detection of warm temperatures, highlighting their importance in behavioural adaptation to innocuous temperatures.

Discussion

Earlier version: Environmental temperatures are detected by sensory nerve fibres innervating the skin. The mechanisms behind warmtemperature detection have recently gained increased attention, with three ion channels, TRPV1, TRPM2 and TRPM8, as the main candidates (7, 8, 11). In this study, we developed a novel temperature preference assay, integrating ambient and floor temperatures, to investigate the roles of TRPV1 and TRPM2 in temperature detection. Our results reveal distinct behavioural responses to warm temperatures mediated by these channels. Applying a modelling framework to the animal’s behaviour, we observed unique deficits in TRPV1 and TRPM2 knockout animals, compared to wildtype mice. On the cellular level, the loss of either TRPV1 or TRPM2 resulted in a decreased proportion of WSNs, with TRPV1 playing a pivotal role in detecting rapid, dynamic temperature changes, while TRPM2-loss affected the population response of WSNs.

Revised version: Environmental temperatures are detected by sensory nerve fibres innervating the skin. The mechanisms behind warm-temperature detection have recently gained increased attention, with three ion channels, TRPV1, TRPM2 and TRPM8, as the main candidates (7, 8, 11). In this study, we developed a novel temperature preference assay, integrating ambient and floor temperatures, to investigate the roles of TRPV1 and TRPM2 in temperature detection. Our results reveal distinct behavioural responses to warm temperatures mediated by these channels. Applying a modelling framework to the animal’s behaviour, we observed unique deficits in TRPV1 and TRPM2 knockout animals, compared to wildtype mice. On the cellular level, the loss of either TRPV1 and, to a lesser extent, TRPM2 resulted in a decreased proportion of WSNs, with TRPV1 playing a pivotal role in detecting rapid, dynamic temperature changes, while TRPM2-loss appeared to affect the population response of WSNs.

Behavioural analysis in temperature preference assays

The introduced chamber preference assay, integrating both ambient and floor temperatures, improves on the conventional temperature preference assays. Notably, at 34 °C, a temperature that is close to the thermoneutral 31 °C, animals demonstrated a clear avoidance of the 34 °C side in the CPT, but failed to do so in the conventional TPP assay (Figure S1I). This preference underscores the importance of integrating multiple sensory inputs — ambient “air” and contact temperatures — in forming a coherent thermal perception, a complexity often overlooked in simpler thermal assays.

Consistent with previous findings, our results reveal that the absence of TRPM2 impedes the development of a preference for warmer temperatures (8). In contrast, animals lacking TRPV1 exhibited behavioural patterns similar to their wild-type counterparts, spending comparable amounts of total time in the warm chamber (Figure 1E-F).

Earlier version: Intriguingly, a finer characterization of the dynamics of the animal behaviour in the assay revealed differences between Trpv1−/− and Trpm2−/− animals, particularly in the frequency of crossings between chambers and the time spent in each chamber (Figure 1G-J). These behavioural nuances were further elucidated by modelling the behaviour with an evidenceaccumulation model (Figure S7 and Figure 5). This model, a novel approach for such behavioural assays in general and for temperature as a sensory modality in particular, uncovered an impaired process of evidence accumulation within the warm chambers in Trpm2−/− animals. Moreover, we could explain the more frequent chamber crossings of Trpv1−/− animals by the fact that they accumulated evidence (information of preferred temperature) more error-prone and thus erratically (Figure 5I).

Revised version: Intriguingly, a finer characterization of the dynamics of the animal behaviour in the assay revealed differences between Trpv1−/− and Trpm2−/− animals, particularly in the frequency of crossings between chambers and the time spent in each chamber (Figure 1G-J). These behavioural nuances were further elucidated by modelling the behaviour with an evidence-accumulation model (Figure S7 and Figure 5). This model, a novel approach for such behavioural assays in general and for temperature as a sensory modality in particular, uncovered an impaired process of evidence accumulation within the warm chambers in Trpm2−/− animals. Moreover, we could explain the more frequent chamber crossings of Trpv1−/− animals by the fact that they accumulated evidence (information of preferred temperature) more error-prone and thus erratically (Figure 5I).

Cellular insights into warm-temperature sensation

Earlier version: Cultures from Trpv1−/− animals exhibited a substantial decrease in the proportion of WSNs (Figure 2). This is similar to previous studies from trigeminal neurons, where the loss of TRPV1 led to a complete loss of warm-temperature responses (7). While the role of TRPV1 was more salient in the warm-temperature range, Trpm2 knockouts displayed a reduction in temperature responsiveness across a broader spectrum, extending into hotter temperatures (Figure 2, Figure S5). This is in line with previous studies highlighting a loss in warm- and heat-responsive cells in Trpm2−/− cells (8–10).

Revised version: Cultures from Trpv1−/− animals exhibited a substantial decrease in the proportion of WSNs (Figure 2). This is similar to previous studies from trigeminal neurons, where the loss of TRPV1 led to a complete loss of warm-temperature responses (7). While the role of TRPV1 was more salient in the warm-temperature range, Trpm2 knockouts displayed a reduction in temperature responsiveness across a broader spectrum, extending into hotter temperatures, albeit the overall effect Trpm2-deletion had on temperature responses in DRG cultures appeared very subtle (Figure 2, Figure S5). This is in line with previous studies highlighting only a subtle loss in warm-/heat-responsiveness in DRG cells of Trpm2−/− animals (8–10).

Earlier version: Interestingly, in initial experiments using DRG neurons from Trpm2−/− and Trpv1−/− animals cultured overnight, we failed to reproduce the previously reported reduction in WSNs in Trpm2−/− and Trpv1−/− sensory neurons (Figure S4H, (7, 8)). The inability to reproduce the aforementioned cellular phenotypes in cultured sensory neurons might be due to two factors: the abundance of WSNs and the variability in their proportions between experiments and animals (Figure 2 and Figure S5). These require a larger sampling of sensory neurons from multiple animals for a reliable estimation of effects, something that is often lacking in previous studies of cellular warm-temperature-detection (8). Furthermore, overnight cultures, which are the de facto standard in the field, might be more akin to an injury model (Figure S4). The three-day cultures presented in this study allow the cells the time to partially re-generate from the harsh dissociation procedure (45), and pose an alternative that more closely resembles the physiological condition.

Revised version: Interestingly, in initial experiments using DRG neurons from Trpm2−/− and Trpv1−/− animals cultured overnight, we failed to reproduce the previously reported reduction in WSNs in Trpm2−/− and Trpv1−/− sensory neurons (Figure S4H, (7, 8)). The inability to reproduce the aforementioned cellular pheno-types in cultured sensory neurons might be due to two factors: the abundance of WSNs and the variability in their proportions between experiments and animals (Figure 2 and Figure S5). These require a larger sampling of sensory neurons from multiple animals for a reliable estimation of effects, something that is often lacking in previous studies of cellular warm-temperature-detection (8). Furthermore, overnight cultures, which are the de facto standard in the field, might be more akin to an injury model (Figure S4). The three-day cultures presented in this study allow the cells the time to partially regenerate from the harsh dissociation procedure (45), and pose an alternative that more closely resembles the physiological condition.

TRPV1: bridging cellular data with behavioural patterns

The cellular data predict that animals lacking TRPV1 would have large deficits in their ability to detect warm temperatures. Yet, overall, Trpv1−/− animals stay in the thermoneutral chamber for a similar proportion of time as wildtype controls. Analysis of the remaining WSNs in Trpv1−/− animals revealed a critical insight: these neurons predominantly respond during the static phase of the temperature stimuli.

Earlier version: In the CPT, animals frequently transition between chambers, experiencing rapid temperature changes upon crossing, but then spend most of their time in an isothermal environment (Figure S1H). This suggests that Trpv1−/− animals primarily rely on static temperature information for thermal detection, rather than rapidly fluctuating temperatures. This is reflected in the higher number of crossings between the different thermal chambers, coupled with shorter visits to the hotter chambers. This set of results led us to hypothesize that rapidly changing thermal information perceived during the transitions of the animals is not properly detected by Trpv1−/− animals. This hypothesis is further supported by the behavioural model (Figure 5). It indicates that, Trpv1−/− animals exhibit a higher drift rate in warmer test chambers, suggesting an avoidance of these temperatures. However, the increased noise in the driftdiffusion model points to a less reliable temperature detection mechanism. This implies that, despite the substantial loss of WSNs, the remaining neuronal population provides sufficient information for the detection of warmer temperatures, albeit with reduced precision.

Revised version: In the CPT, animals frequently transition between chambers, experiencing rapid temperature changes upon crossing, but then spend most of their time in an isothermal environment (Figure S1H). This suggests that Trpv1−/− animals primarily rely on static temperature information for thermal detection, rather than rapidly fluctuating temperatures. This is reflected in the higher number of crossings between the different thermal chambers, coupled with shorter visits to the hotter chambers. This set of results led us to hypothesize that rapidly changing thermal information perceived during the transitions of the animals is not properly detected by Trpv1−/− animals.

Added in the revised version: This hypothesis is further supported by the behavioural model (Figure 5). It indicates that, Trpv1−/− animals exhibit a higher drift rate in warmer test chambers, suggesting an avoidance of these temperatures. However, the increased noise in the drift-diffusion model points to a less reliable temperature detection mechanism. Although noise in drift diffusion models can encompass various sources of variability—ranging from peripheral sensory processing to central mechanisms like attention or motor initiation—the most parsimonious interpretation in our study aligns with a perceptual deficit, given the altered temperature-responsive neuronal populations we observed. This implies that, despite the substantial loss of WSNs, the remaining neuronal population provides sufficient information for the detection of warmer temperatures, albeit with reduced precision.

The reduced precision might stem from the loss of dynamic temperature responders (Figure 3). These WSNs might be crucial in detecting a rapid change of temperature (e.g. when the animals move across different thermal environments). This is highlighted by findings from TRPV1-overexpressing animals. These animals, equipped with an enhanced ability to respond to dynamic temperature changes (Figure 4), have a higher drift rate and lower noise levels in warmer chambers in the model (Figure 5). These characteristics lead to a faster and more precise choice in the CPT. Collectively, these results highlight the direct role of TRPV1, and its expression levels, in the precise temporal detection of warm temperatures. This could also explain the consistent ––albeit subtle–– involvement of TRPV1 in operant assays of temperature perception where temperature stimuli are applied rapidly (7, 11).

TRPM2: cellular mechanisms and behavioural implications

Earlier version: A reversed scenario unfolds for TRPM2. The behavioural data suggests a strong deficit in detecting warm temperatures. Trpm2−/− WSNs, are less abundant compared to WSNs from wildtype DRG cultures, albeit to a lesser degree compared to Trpv1−/− WSNs (Figure 2). Trpm2−/− WSNs also do not differ from wildtype WSNs in their response timings (Figure 3). It is tempting to hypothesize, given the results from Trpv1−/−-cultures, that TRPM2 affects the static phase of temperature detection. We conducted various analyses on the static responses of Trpm2−/− WSNs (data not shown), but failed to uncover significant differences to wildtype WSNs.

Revised version: A reversed scenario unfolds for TRPM2. The behavioural data suggests a strong deficit in detecting warm temperatures. In Trpm2−/− cultures, WSNs appear less abundant compared to wildtype DRG cultures, albeit the contribution of Trpm2 appears to be less robust compared to that of Trpv1 (Figure 2). Trpm2−/− WSNs also do not differ from wildtype WSNs in their response timings (Figure 3). It is tempting to hypothesize, given the results from Trpv1−/−-cultures, that TRPM2 affects the static phase of temperature detection. We conducted various analyses on the static responses of Trpm2−/− WSNs (data not shown), but failed to uncover significant differences to wildtype WSNs.

Earlier version: The only difference we found when comparing Trpm2−/− cultures to WSNs from wildtype and Trpv1−/− animals was their increased response magnitude to the first stimulus, without affecting the subsequent stimuli (Figure 2). Previous studies hypothesized that peripheral warm and hot temperature perception requires a combination of population and graded rate coding (7, 21). This means that a series of temperature stimuli with increasing intensity should activate and recruit more neurons, as well as increase their response magnitude. Trpm2−/− WSNs deviate from this model by having a U-shaped response magnitude to increasing temperature stimuli (Figure 2F). Two different temperatures leading to similar response magnitudes could impair the animals ability to differentiate these two temperatures and thereby lead to the observed phenotypes in the CPT (Figure 1 and 5). The question remains as to why the loss of TRPM2 would lead to an increase in the response magnitude. Given the relatively minor cellular phenotype observed by TRPM2-deletion, there are several alternative hypothesis that could explain the discrepancy between the behavioral and cellular phenotypes:

Revised version: The only difference we found when comparing Trpm2−/− cultures to WSNs from wildtype and Trpv1−/− animals was their increased response magnitude to the first stimulus, without affecting the subsequent stimuli (Figure 2). Previous studies hypothesized that peripheral warm and hot temperature perception requires a combination of population and graded rate coding (7, 21). This means that a series of temperature stimuli with increasing intensity should activate and recruit more neurons, as well as increase their response magnitude. Trpm2−/− WSNs deviate from this model by having a U-shaped response magnitude to increasing temperature stimuli (Figure 2F). Two different temperatures leading to similar response magnitudes could impair the animals ability to differentiate these two temperatures and thereby lead to the observed phenotypes in the CPT (Figure 1 and 5). The question remains as to why the loss of TRPM2 would lead to an increase in the response magnitude, and whether this relatively small effect would have any bearing on temperature sensing in vivo.

Added in the revised version: Given the relatively minor cellular phenotype observed by TRPM2-deletion, there are several alternative hypotheses that could explain the discrepancy between the behavioral and cellular phenotypes:

Earlier version: It is possible that permanent deletion of TRPM2 in the Trpm2 KO mice results in developmental defects and/or compensatory mechanisms that mask a more prominent phenotype that might occur if the channel would be acutley blocked. In the absence of specific TRPM2 antagonists such an experiment is currently not possible. Another explanation involves the specific population of WSNs reliant on TRPM2. Single-cell sequencing and functional analyses of dorsal root ganglia suggest that WSNs form two, genetically distinct populations, one of which prominently expresses Trpm2 (46, 47). Different genetic/functional populations of sensory neurons often show diverging spinal innervation, with different upstream processing pathways (47–50). The genetic separation could hint at multiple neural innervation routes for innocuous temperature information. In this context, the loss of TRPM2 might specifically impair warm-temperature perception in thermoregulation-specific innervation pathways, without significantly affecting perceptual temperature discrimination performance (11).

Revised version: It is possible that permanent deletion of TRPM2 in the Trpm2−/− mice results in developmental defects and/or compensatory mechanisms that mask a more prominent phenotype that might occur if the channel would be acutley blocked. In the absence of specific TRPM2 antagonists such an experiment is currently not possible. Another explanation involves the specific population of WSNs reliant on TRPM2. Single-cell sequencing and functional analyses of dorsal root ganglia suggest that WSNs form two, genetically distinct populations, one of which prominently expresses Trpm2 mRNA transcripts (46, 47). Different genetic/functional populations of sensory neurons often show diverging spinal innervation, with different upstream processing pathways (47–50). The genetic separation could hint at multiple neural innervation routes for innocuous temperature information. In this context, the loss of TRPM2 might specifically impair warm-temperature perception in thermoregulation-specific innervation pathways, without significantly affecting perceptual temperature discrimination performance (11).

Earlier version: A third possibility implies that the role of TRPM2 in temperature detection might extend beyond its expression in sensory neurons. Particularly, it is possible that TRPM2 mediates part of its effect via the hypothalamic preoptic area (POA), a temperature sensitive brain region involved in body temperature regulation (51). Preoptic TRPM2 has been shown to mediate autonomic thermoregulatory responses upon warm-temperature stimulation (52–54). Preoptic temperature pathways may not only drive autonomic thermoregulatory responses, but can also influence temperature preference behaviour (51, 55). Considering that, within the time frame of the chamber preference test, ambient temperature changes are directly transferred to the POA, it is possible – but not yet tested – that preoptic TRPM2 is involved in the choice of comfort-temperature. The use of conditional Trpm2 knock out animals may help to clarify this aspect in future studies. On a more general note and considering the choice of the behavioral assay used in different studies, either an operant task or a preference assay without training of the mice: depending on the task the animals perform, likely requires qualitatively and quantitatively different thermal inputs, thereby possibly explaining the different phenotypes observed in individual TRP channel knockout mouse models.

Revised version: A likely third possibility implies that the role of TRPM2 in temperature detection extends beyond its expression in sensory neurons. Particularly, it is possible that TRPM2 mediates part of its effect via the hypothalamic preoptic area (POA), a temperature sensitive brain region involved in body temperature regulation (51). Preoptic TRPM2 has been shown to mediate autonomic thermoregulatory responses upon warm-temperature stimulation (52–54). Preoptic temperature path-ways may not only drive autonomic thermoregulatory responses, but can also influence temperature preference behaviour (51, 55). Considering that, within the time frame of the chamber preference test, ambient temperature changes are directly transferred to the POA, it is possible – but not yet tested – that preoptic TRPM2 is involved in the choice of comfort-temperature. The use of conditional Trpm2 knock out animals may help to clarify this aspect in future studies.

Added in the revised version: On a more general note and considering the choice of the behavioral assay used in different studies, either an operant task or a preference assay without training of the mice: depending on the task the animals perform, likely requires qualitatively and quantitatively different thermal inputs, thereby possibly explaining the different phenotypes observed in individual TRP channel knockout mouse models.

Implications of findings

This study introduces an alternative protocol to culture DRG neurons to reduce their (thermal) hypersensitivity, an innovative behavioural assay,, and methodologies for analysing animal behaviour in temperature preference assays. We emphasize the importance of examining the dynamics of perceptual decision-making and incorporating behavioural modelling. Significantly, we demonstrate that TRPV1 and TRPM2 channels contribute differently to temperature detection, supported by behavioural and cellular data. This research not only advances our understanding of thermal perception mechanisms but also adds new dimensions for integrating cellular and behavioural data to study the neural foundations of temperature sensation.

Limitations of the study

Earlier version: In this study, only male animals were used to study temperature preference in the CPT assay. It is possible that sex differences modulate thermal preference in a TRP channel-dependent manner, similar to what has been shown for an operant-based thermal choice assay (11, 56). A systematic comparison of TRP-channel involvement in sex differences in thermal detection awaits future analysis. The 3day DRG cultures used in this study display a response profile that is more similar to the (rare) warmth responses detected in vivo compared to acute DRG cultures (which display exaggerated/sensitized warmth/heat responses). However, longer culturing times may result in cellular changes and a drift of neuronal identity away from their native state. Future in vivo DRG recording/imaging studies using Trp knock out mouse models will help to reveal how warmth responses are coded in the native cellular setting.

Revised version: In this study, only male animals were used to study temperature preference in the CPT assay. It is possible that sex differences modulate thermal preference in a TRP channel-dependent manner, as previously suggested (56). A systematic comparison of TRP-channel involvement in sex differences in thermal detection across different thermal modalities (operant-based choice assays, thermal preference chamber test and others) awaits future analysis. The 3-day DRG cultures used in this study display a response profile that is more similar to the (rare) warmth responses detected in vivo compared to acute DRG cultures (which display exaggerated/sensitized warmth/heat responses). However, longer culturing times may result in cellular changes and a drift of neuronal identity away from their native state. Future in vivo DRG recording/imaging studies using TRP-channel knock out mouse models will help to reveal how warmth responses are coded in the native cellular setting.

Added in the revised version: Methods

Added in the revised version: Animals and housing

Added in the revised version: All animal care and experimental procedures were approved by the local council (Regierungspräsidium Karlsruhe, Germany) under protocol numbers G-201/16 and T05-19. Animals were kept under specific-pathogen-free (SPF) conditions and a 12-hour day-night cycle. Housing temperature and humidity were maintained at 22 ± 2 °C and 50 % to 60 %, respectively. Animals were fed ad-libitum with Altromin Rod 16 or Rod 18 animal food. The housing environment was enriched using Crincklets Nest-Pads and ABBEDD LT-E-001 bedding. For this study, only male animals were used, as we aimed to compare our results with previous studies which exclusively used male animals (7, 8). Mice between 6 and 25 weeks of age were used for the experiments.

Added in the revised version: Thermal preference chamber design, operation, and video capture

Added in the revised version: The thermal preference chamber consisted of two expanded polystyrene boxes connected using plastic glue and sealed with silicone (Figure S1A) with dimensions of 26.1 cm x 60.3 cm x 19.8 cm (w x l x h). The 4.3 cm-thick Styrofoam walls provided the necessary thermal insulation for the experiments. Inside the enclosure, the animals’ movement was limited by a 13 cm x 31.4 cm x 15.5 cm steel cage placed on top of a stainlesssteel baseplate. To create two thermally isolated chambers, the enclosure, cage, and baseplate were adjusted to form a 4.3 cm x 6.3 cm x 6.3 cm tunnel. The baseplate and cage were custom-built by our institute’s mechanical workshop. The cover combined foam and wood as insulators, with an acrylic glass inset for observation purposes.

Added in the revised version: Temperature within the chamber was regulated by two Peltier elements attached to heat sinks, each connected to a generic computer fan for efficient temperature distribution. To avoid overheating, the Peltier elements were connected to a Multitemp III circulating water pump (Amersham Biosciences) set to 28 °C. These elements were managed by a modified Siemens LOGO TD! controller, programmed for precise temperature adjustments, and accessed using the LOGO! Soft Comfort (Siemens) software. For monitoring, two Physitemp IT-18 flexible thermocouples were attached to the chamber walls, serving as reference thermometers. Data capture was conducted using a customer-grade webcam (Spedal), linked to the UCLA miniscope project’s capture software, operating at 20 Hz to 30 Hz. The output files from each recording session were concatenated using ffmpeg software for subsequent processing. For control experiments using the two-plate temperature preference test (Figure S1, we used the BIO-T2CT system by BioSeb with the same camera setup as described above.

Added in the revised version: Thermal preference tests, processing, and analysis

Added in the revised version: Animals were transported to and acclimatized in the experimental room for at least 24 h before the experiments. The room maintained a dim light setting and a 12-hour day-night cycle. To ensure temperature stability, the setup was allowed to stabilize for 90 min before starting an experiment (Figure S1B-D). Once the setup reached a stable temperature, the lid was briefly opened (Figure S1), and an animal was placed into the enclosure, the lid closed again, and allowed to roam freely for at least 30 min, before replacing it with the next animal.

Added in the revised version: Due to the dim light conditions and the reflections on the cover of the enclosure (see Figure 1 for an example), conventional animal tracking approaches that rely on image contrast failed to provide robust outputs (data not shown). Therefore, we employed a neural-network based approach, namely DeepLabCut, for tracking the animals in the setup (57). For this purpose, we trained a ResNet50 network to track specific points on the animal: the snout, right ear, left ear, body centre, tail base, and tail tip (Figure 1). Additionally, we tracked eight reference points (four cage corners and four cage tops) in the cage for normalization purposes. The same model was also trained on reference frames acquired using the BIO-T2CT system. The resulting model generalized well to both setups, ensuring comparability of the outputs.

Added in the revised version: The DeepLabCut predictions were then further curated by replacing bad predictions (< 95% likelihood) with missing values, and removing recording sessions where > 25% were missing (usually due to the animal escaping the inner cage). Furthermore, we removed sessions in which the animals showed excessive climbing behaviour (> 95% of time between upper and lower cage corners). For the remaining sessions, the missing values were linearly interpolated and the centroid of the ears, body centre, and tail-base was used as the position of the animals. The X- and Y-positions of the centroid were then scaled to the tracked cage corner points to correct for minor movements of the camera or the setup. The resulting X- and Y-position time courses were then downsampled to 1 Hz. Only the first 30 min of the recording were kept. Shorter sessions were removed.

Added in the revised version: Primary sensory neuron culture

Added in the revised version: Adult primary DRG cultures were prepared from 6-15 week old animals as described previously (36). Briefly, the animals were culled via isoflurane overdose, and their spinal columns excised and separated from muscle tissue. The spinal column was then cut lengthwise and the DRGs collected, freed from nerve branches, halved, and treated with a collagenase solution (1.25 mg/mL in Ringer’s solution) for 1 h at 37 °C with gentle inversion every 15 min. This was followed by a 15 min trypsin digestion (2.5 mg/mL) at 37 °C, repeated trituration and suspension in complete culturing medium (DMEM/F12 w/o Glutamin,10% heat-inactivated FCS, 2 mM L-glutamine (Glutamax), 1x Anti-biotic/mitotic), and centrifugation at 900 rpm for 10 min over a BSA solution (150 mg/mL) to pellet the cells. The supernatant was discarded, the pellet was resuspended in culturing media, and spotted onto PDL- and Laminin-coated glass coverslips (5 mm). The cells were then left to settle onto the coverslip in the incubator for 1 h at 37 °C, and then covered with culturing media. Cultures were either used the following day (overnight) or kept for three days, with a medium change after the first day.

Added in the revised version: Calcium imaging recordings

Added in the revised version: For calcium imaging, cells cultured on coverslips were incubated for either one or three days. Before imaging, cells underwent a washing process with Ringer’s solution (140 mM NaCl, 5 mM KCl, 2 mM MgCl2, 2 mM CaCl2, 10 mM glucose, and 10 mM HEPES, adjusted to pH 7.4.), followed by loading with the calcium-sensitive dye Cal520-AM (10 μM) and Pluronic acid F-127 (0.05 %) in Ringer’s solution. The cells were incubated for 1 h at 37 °C, then the dye solution was replaced with Ringer’s solution for a further 30 min at room temperature, minimizing light exposure.

Added in the revised version: The perfusion system, a ValveBank II (Automate) to control multiple inflows, and an air200 aquarium pump (Eheim) for outflow, was set to a maximum flow rate of 3 mL/min, facilitating laminar flow in the imaging chamber (RC-22, Warner Instruments). The coverslip was placed near the outlet to minimize movement artefacts. An IT-18 thermocouple (physitemp) was placed close to the coverslip to record the chamber temperature. Imaging settings varied based on the camera used: for CoolSnapHQ2 (Photometrics), exposure was set at 5 ms with 3x gain and 3×3 binning, whereas for Zyla 4.2 (Andor Technology), it was 80 ms with 2×2 binning. We used a Lambda DG-4 as a light source, maintained at 30 % intensity to reduce bleaching.

Added in the revised version: Images were captured using MetaFluor software at 4 Hz or 10 Hz frequencies. Standard experiments involved a 1 min baseline, 25 s stimuli, followed by a 3 min to 5 min recovery period with room-temperature Ringer’s solution (Figures Figure 2 and Figure S5). To identify TRPV1-positive neurons, we used the agonist Capsaicin (1 μM, (3)). A Ringer’s solution with a high potassium concentration (100 mM KCl) was used as a final stimulus to identify neurons. Solutions were heated via glass coils connected to a heated water bath. Each FOV was imaged for a maximum of 60 min, and the usage of coverslips was limited to 2 h post-loading with the dye.

Added in the revised version: Calcium imaging preprocessing and analysis

Added in the revised version: Calcium imaging data were motion-corrected and pre-processed using Suite2p (58). Cell regions of interest (ROIs) were identified using the Cellpose package integrated into Suite2p (59). The mean fluorescence of cells and surrounding neuropil was calculated by Suite2p, and neuropil contamination was corrected by subtracting 70 % of the background neuropil traces from each cell’s fluorescence trace. The corrected data was then imported into a custom R-package, neuroimgr, for further analysis in R (https://github.com/hummuscience/neuroimgr).

Added in the revised version: Normalization was performed using the ΔF/F0 method, where baseline fluorescence (F0) is calculated as the mean fluorescence of the baseline, and ΔF is the change in fluorescence over time. For heatmaps, F0 was estimated using the first 10 s of the experiments. For individual stimuli, the mean of the first 10 frames was used as F0. Heatmaps were generated using the ComplexHeatmap R package. Cells were sorted by the earliest time point where they cross 10 % of their cumulative ΔF/F0 in a given FOV and clustered using the Ward D2 algorithm. ΔF/F0 values smaller than the 0.1 and larger than the 99.9 percentile were clipped.

Added in the revised version: Due to the temperature-sensitivity and loading variability of the calcium dye, a threshold-based approach failed to reliably identify responding cells across experimental days and FOVs (60). Therefore, we used time series classification to identify temperature-responsive cells. For this, calcium traces for each cell and stimulus were normalized, downsampled to 4 Hz, and a sample of 1000 traces across stimuli was manually labelled to create a training dataset. Examples of responsive and non-responsive cells as well as the average of each label are shown in Figure S4A and B. This ground-truth dataset was used to evaluate multiple time series classification algorithms, with MINIROCKET (as implemented in the sktime Python package, (61)) yielding the best results (classification results shown in Figure S4B). The trained classifier was then applied to the remaining cells to identify temperature-responsive cells. A similar approach was used to identify capsaicin-responsive cells. Cells that did not respond to any of the applied stimuli were excluded from the analysis.

Added in the revised version: Drift-diffusion model

Added in the revised version: We employed a Drift-Diffusion Model (DDM) to analyse the behaviour of mice in thermal chamber experiments. The DDM was preferred over simpler models like the Markov switching model, as the latter did not provide satisfactory fits to our data (data not shown). Each parameter in the model can either be fit as a predictor (dependent on genotype, temperature combination, and temperature), fit to the entire data (floating), or fixed to a certain value. A typical drift diffusion model as described in (44) can drift to the upper or lower bound (representing two choices), but our experimental design only offers one choice (to leave the chamber). To reduce the probability of reaching the lower bound, and thereby improve the fit, we fixed the starting point/bias (z) to 0.9. This ensures that the evidence accumulation starts at a point that is much closer to the upper bound (a) than to the lower bound (−a). To choose the best combination of parameters that fits the data, we fit all the data to all remaining combinations of v, sv, and a (Figure S7A and B) and compared them via the expected log pointwise predictive density (ELPD) by Pareto smoothed importance sampling leave-one-out cross-validation (LOO) (62). Unreliable models as per (62) were discarded. The chosen model was constructed to account for variations in both noise (sv) and drift rate (v) for each genotype, temperature comparison, and chamber/temperature (Figure S7A and B, and Equation 1). To accommodate individual differences among animals, we introduced a random effect for each animal in the model. This approach enabled us to capture the unique behavioural patterns of each subject while assessing the general trends across the population.

Added in the revised version: To fit the model, we applied a hierarchical Markov Chain Monte Carlo (MCMC) sampling approach as implemented in the HSSM python package. For the implementation of the hierarchical MCMC, we utilized the No-U-Turn Sampler (NUTS) as implemented in NumPyro, a robust algorithm for efficiently sampling from high-dimensional probability distributions. The tuning phase for all fit models involved 2000 samples, ensuring adequate exploration of the parameter space and helping to achieve convergence. The final model was run with four chains, each drawing 2000 samples.

Added in the revised version: Statistical methods

Added in the revised version: Statistical analyses were conducted using R software. For time course experiments involving repeated measures, Two-way analysis of variance (ANOVA) with repeated measures as implemented in the afex package was conducted. Mauchly’s test was applied to assess the assumption of sphericity, and corrections for violations were made using the Geisser-Greenhouse correction. For comparisons with imbalanced observations or missing data, we fit a linear mixed effect model to the data using the lmer package. In cases of significant outcomes, post-hoc comparisons were performed using estimated marginal means (EMMs) with pairwise contrasts comparing treatments to control groups (facilitated by the emmeans package), with false-discovery rate (FDR) for multiple comparison corrections. For non-parametric data, we applied a Wilcoxon Rank-Sum test as implemented in the rstatix package, coupled with FDR for multiple comparison correction. To assess differences in crossing behaviour, a Cox proportional hazard model was used, as implemented in the survival package. For visit length comparisons, a mixed-effects model was fit using the lme4 package, allowing random effects for animal subjects where appropriate and correcting for the effect of time. Multiple comparisons for the mixed-effects and Cox models were accounted for using the FDR approach within the emmeans and multcomp packages. Only statistically significant results (p <0.05) are shown.